---
title       : Insert the chapter title here2
description : Insert the chapter description here
---


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:8165c5ccf4
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:eff805b059
## <<<New Exercise>>>


*** =video_link
//player.vimeo.com/video/154783078
